```python

```

<img align="center" src="images/course.png" width="800">

# 16720 (B)  3D Reconstruction - Assignment 5 - q3
    Instructor: Kris                          TAs: Arka, Jinkun, Rawal, Rohan, Sheng-Yu

```python
# Helper functions for this assignment. DO NOT MODIFY!!!
"""
Helper functions.
"""

import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import scipy
import cv2
import nbimporter
from q1 import eightpoint, sevenpoint, camera2, _epipoles, calc_epi_error, toHomogenous, _singularize
from q2 import find_M2, plot_3D

def plot_3D_dual(P_before, P_after):
    matplotlib.use('TkAgg')
    fig = plt.figure()
    ax = fig.add_subplot(111, projection='3d')
    ax.set_title("Blue: before; red: after")
    ax.scatter(P_before[:,0], P_before[:,1], P_before[:,2], c = 'blue')
    ax.scatter(P_after[:,0], P_after[:,1], P_after[:,2], c='red')
    while True:
        x, y = plt.ginput(1, mouse_stop=2)[0]
        plt.draw()


```

## Q3: Bundle Adjustment
Bundle Adjustment is commonly used as the last step of every feature-based 3D reconstruction algorithm. Given a set of images depicting a number of 3D points from different viewpoints, bundle adjustment is the process of simultaneously refining the 3D coordinates along with the camera parameters. It minimizes reprojection error, which is the squared sum of distances between image points and predicted points. In this section, you will implement bundle adjustment algorithm by yourself. Specifically,


- In Q3.1, you need to implement a RANSAC algorithm to estimate the fundamental matrix F and all the inliers.
- In Q3.2, you will need to write code to parameterize Rotation matrix $\mathbf{R}$ using [Rodrigues formula](https://en.wikipedia.org/wiki/Rodrigues\%27\_formul) (Please check [this pdf](https://www2.cs.duke.edu/courses/fall13/compsci527/notes/rodrigues.pdf) for a detailed explanation), which will enable the joint optimization process for Bundle Adjustment.
- Q3.3, you will need to first write down the objective function in rodriguesResidual, and do the bundleAdjustment.

### Q3.1 RANSAC for Fundamental Matrix Recovery (15 pt implementation)

In some real world applications, manually determining correspondences is infeasible and often there will be noisy correspondences. Fortunately, the RANSAC method seen (and implemented in previous assignments) in class can be applied to the problem of fundamental matrix estimation.

Implement the above algorithm with the signature:
```
[F, inliers] = ransacF(pts1, pts2, M)
```

where `M` is defined in the same way as when we calculate the fundamental matrix and inliers is a boolean vector of size equivalent to the number of points. Here inliers are set to true only for the points that satisfy the threshold defined for the given fundamental matrix F.

We have provided some noisy coorespondances in some\_corresp\_noisy.npz in which around $75\%$ of the points are inliers. Compare the result of RANSAC with the result of the eight-point algorithm when ran on the noisy correspondences. 

**Hints:** Use the seven point to compute the fundamental matrix from the minimal set of points. Then compute the inliers, and refine your estimate using all the inliers.


<!-- grade_id: cell-450fffa620eaf4bc, points: 0, grade: false, solution: true -->
```python
def ransacF(pts1, pts2, M):
    '''
    Q3.1: RANSAC method.
        Input:  pts1, Nx2 Matrix
                pts2, Nx2 Matrix
                M, a scaler parameter
        Output: F, the fundamental matrix
                inlier_curr, Nx1 bool vector set to true for inliers
    ***
    Hints:
    (1) You can use the calc_epi_error from q1 with threshold to calcualte inliers. Tune the threshold based on 
        the results/expected number of inliners. You can also define your own metric. 
    (2) Use the seven point alogrithm to estimate the fundamental matrix as done in q1
    (3) Choose the resulting F that has the most number of inliers
    (4) You can increase the nIters to bigger/smaller values
        
    '''
    N = pts1.shape[0]
    pts1_homo, pts2_homo = toHomogenous(pts1), toHomogenous(pts2)
    threshold = 10
    max_itaration = 1000
    best_inlier = 0
    inlier_curr = None
    # ----- TODO -----
    # YOUR CODE HERE

    # T = np.diag([1/M,1/M,1])
    # pts1_homo_norm = np.matmul(pts1_homo, T)
    # pts2_homo_norm = np.matmul(pts2_homo, T)
    for i in range(max_itaration):
        
        idxs = np.random.choice(N, size=7, replace=False)
        # coords1 = np.array([pts1_homo[idxs,0], pts1_homo[idxs,1]])
        coords1 = np.array(pts1[idxs])
        # coords2 = np.array([pts2_homo[idxs,0], pts2_homo[idxs,1]])
        coords2 = np.array(pts2[idxs])
        # print(coords1.shape, coords2.shape)
        F = sevenpoint(coords1, coords2, M)

        for f in F:
            dist = calc_epi_error(pts1_homo, pts2_homo, f)
            inlier_curr = dist < threshold
            inlier_count = np.sum(inlier_curr)
            if inlier_count > best_inlier:
                best_inlier = inlier_count
                best_inlier_curr = inlier_curr
                F_best = f
    F = F_best
    inlier_curr = best_inlier_curr
    # raise NotImplementedError()
    return F, inlier_curr


np.random.seed(0)
np.set_printoptions(precision=4, suppress=1)
correspondence = np.load('data/some_corresp_noisy.npz') # Loading noisy correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1 = plt.imread('data/im1.png')
im2 = plt.imread('data/im2.png')
F, inliners = ransacF(pts1, pts2, M=np.max([*im1.shape, *im2.shape]))
# print(len(F))

```

<!-- grade_id: q3_1_a, points: 15, grade: true, solution: false -->
```python
# Full set of tests; you will get full points for coding if passing the following tests. 
try:
    assert(np.sum(inliners) > len(pts1) * 0.7)
    print('Test passed!')
except:
    raise AssertionError('Test failed')
```

Output:
```
Test passed!
```

### Q3.2 Rodrigues and Invsere Rodrigues (10 pt implementation)
So far we have independently solved for the camera matrix, $\mathbf{M}_j$ and 3D projections, $\textbf{w}_i$. In bundle adjustment, we will jointly optimize the reprojection error with respect to the points $\textbf{w}_i$ and the camera matrix $\textbf{w}_j$.

$$
err = \sum_{ij} ||\textbf{x}_{ij} - Proj(\mathbf{C}_j, \textbf{w}_i)||^2,
$$
where $\textbf{w}_j = \mathbf{K}_j \mathbf{M}_j$.

For this homework, we are going to only look at optimizing the extrinsic matrix. The rotation matrix forms the Lie Group $\textbf{SO}(3)$ that doesn't satisfy the addition operation so it cannot be directly optimized. Instead, we parameterize the rotation matrix to axis angle using Rodrigues formula to the Lie Algebra $\mathfrak{so}(3)$, which is defined in $\mathbb{R}^3$. through which the least squares optimization process can be done to optimize the axis angle. Try to implement function

```
R = rodrigues(r)
```

as well as the inverse function that converts a rotation matrix $\mathbf{R}$ to a Rodrigues vector $\mathbf{r}$

```
r = invRodrigues(R)
```

Please refer to [Rodrigues formula](https://en.wikipedia.org/wiki/Rodrigues\%27\_formul)  and [this pdf](https://www2.cs.duke.edu/courses/fall13/compsci527/notes/rodrigues.pdf) for reference.


<!-- grade_id: cell-f7107c598b84fe13, points: 0, grade: true, solution: true -->
```python
def rodrigues(r):
    '''
    Q3.2: Rodrigues formula.
        Input:  r, a 3x1 vector
        Output: R, a 3x3 rotation matrix
    '''
    # ----- TODO -----
    # YOUR CODE HERE\
    r = r.reshape((3,1))
    theta = np.linalg.norm(r)
    I = np.diag(np.ones(3))
    if theta == 0:
        return I
    u = r/theta
    # print("u is",u)
    # print("u[0] is",u[0,0])
    u_x = np.array([[0, -u[2,0], u[1,0]],
                    [u[2,0], 0 , -u[0,0]],
                    [-u[1,0], u[0,0], 0]])
    R = I*np.cos(theta) + (1-np.cos(theta))*u@u.T + u_x*np.sin(theta)
    # raise NotImplementedError()
    return R


def invRodrigues(R):
    '''
    Q5.2: Inverse Rodrigues formula.
        Input:  R, a 3x3 rotation matrix
        Output: r, a 3x1 vector
    '''
    # ----- TODO -----
    # YOUR CODE HERE

    I = np.diag(np.ones(3))
    A = (R - R.T)/2
    rho = np.array([[A[2,1]], [A[0,2]], [A[1,0]]])
    s = np.linalg.norm(rho)
    c = (R.trace() - 1) / 2
    
    if s == 0 and c == 1:
        r = np.zeros((3,1))
        r = r.reshape(3)
        return r
    
    elif s == 0 and c == -1:
        mat = R+I
        v = mat[:, mat.any(axis=1)][0]
        u = v / np.linalg.norm(v)
        r = u * np.pi
        if np.linalg.norm(r) == np.pi and ((r[0] == 0 and r[1] == 0 and r[2] < 0) or (r[0] == 0 and r[1] < 0) or (r[0] < 0)):
            r = -r
            r = r.reshape(3)
            return r
        else:
            r = r.reshape(3)
            return r
    
    else:
        u = rho/s
        theta = np.arctan2(s, c)
        r = u * theta
        r = r.reshape(3)
        return r


    # raise NotImplementedError()
    # return r

```

<!-- grade_id: q3_2_a, points: 3, grade: true, solution: false -->
```python
# Simple Tests to verify your implmentation:
from scipy.spatial.transform import Rotation as sRot
rotVec = sRot.random()
mat = rodrigues(rotVec.as_rotvec())
print(mat)


try:
    assert(np.linalg.norm(rotVec.as_rotvec() - invRodrigues(mat)) < 1e-3)
    assert(np.linalg.norm(rotVec.as_matrix() - mat) < 1e-3)
    print('Test passed!')
except:
    print('Test failed!')

```

Output:
```
[[ 0.9106  0.027   0.4125]
 [ 0.3733 -0.4822 -0.7926]
 [ 0.1774  0.8757 -0.4491]]
Test passed!
```

<!-- grade_id: q3_2_b, points: 7, grade: true, solution: false -->
```python
# Hidden Tests
```

### Q3.3 Bundle Adjustment (10 pt writeup)

In this section, you need to implement the bundle adjustment algorithm. Using the parameterization you implemented in the last question, write an objective function for the extrinsic optimization:

```
residuals = rodriguesResidual(K1, M1, p1, K2, p2, x)
```
where x is the flattened concatenation of $\mathbf{w}$, $\mathbf{r}_2$, and $\mathbf{t}_2$.
$\mathbf{w}$ are the 3D points; $\mathbf{r}_2$ and $\mathbf{t}_2$ are the rotation (in the Rodrigues vector form) and translation vectors associated with the projection matrix $\mathbf{M}_2$; $p1$ and $p2$ are 2D coordinates of points in image 1 and 2, respectively. The `residuals` are the difference between the original image projections and the estimated projections (the square of $2$-norm of this vector corresponds to the error we computed in Q3.2):
```
residuals = numpy.concatenate([(p1-p1').reshape([-1]), (p2-p2').reshape([-1])])
```

Use this objective function and Scipy's nonlinear least squares optimizer $\texttt{leastsq}$ write a function to optimize for the best extrinsic matrix and 3D points using the inlier correspondences from some_corresp_noisy.npz and the RANSAC estimate of the extrinsics and 3D points as an initialization.

```
[M2, w, o1, o2] = bundleAdjustment(K1, M1, p1, K2, M2_init, p2, p_init)
```

Try to extract the rotation and translation from `M2_init`, then use `invRodrigues` you implemented previously to transform the rotation, concatenate it with translation and the 3D points, then the concatenate vector are variables to be optimized. After obtaining optimized vector, decompose it back to rotation using `Rodrigues` you implemented previously, translation and 3D points coordinates.

<span style='color:red'>**Output:**</span> In your write-up: include an image of output of the `plot_3D_dual` function by passing in the original 3D points and the optimized points. Also include the before and after reprojection error for the `rodriguesResidual` function.



<!-- grade_id: cell-21b599d494aeefec, points: 0, grade: false, solution: true -->
```python
def get_residual_norm(x, K1, M1, p1, K2, p2):
    residuals = rodriguesResidual(K1, M1, p1, K2, p2,x )
    return np.linalg.norm(residuals)


def rodriguesResidual(K1, M1, p1, K2, p2, x):
    '''
    Q3.3: Rodrigues residual.
        Input:  K1, the intrinsics of camera 1
                M1, the extrinsics of camera 1
                p1, the 2D coordinates of points in image 1
                K2, the intrinsics of camera 2
                p2, the 2D coordinates of points in image 2
                x, the flattened concatenationg of P, r2, and t2.
        Output: residuals, 4N x 1 vector, the difference between original 
                and estimated projections
    '''
    residuals = None
    # ----- TODO -----
    # YOUR CODE HERE
    
    N = p1.shape[0]
    t = x[-3:]
    t = t.reshape((3,1))
    r = x[-6:-3]
    r = r.reshape((3,1))
    R = rodrigues(r)

    M2 = np.hstack((R,t))

    X = x[:3*N]
    X = X.reshape((N,3))
    # X_homo = toHomogenous(X)
    # print(X.shape)
    X_homo = np.concatenate((X, np.ones((X.shape[0],1))), axis=1)
    # print(X_homo.shape)
    P1 = K1 @ M1
    x1_reproj = P1 @ X_homo.T
    x1_reproj = x1_reproj/x1_reproj[-1]
    # print("before t",x1_reproj.shape)
    x1_reproj = np.transpose(x1_reproj)
    # print("after t",x1_reproj[:,:2].shape)

    P2 = K2 @ M2
    x2_reproj = P2 @ X_homo.T
    x2_reproj = x2_reproj/x2_reproj[-1]
    x2_reproj = np.transpose(x2_reproj)
    # print(p1.shape)
    residuals = np.concatenate([(p1-x1_reproj[:,:2]).reshape([-1]), (p2-x2_reproj[:,:2]).reshape([-1])])
    

    # raise NotImplementedError()
    return residuals


def bundleAdjustment(K1, M1, p1, K2, M2_init, p2, P_init):
    '''
    Q3.3 Bundle adjustment.
        Input:  K1, the intrinsics of camera 1
                M1, the extrinsics of camera 1
                p1, the 2D coordinates of points in image 1
                K2,  the intrinsics of camera 2
                M2_init, the initial extrinsics of camera 1
                p2, the 2D coordinates of points in image 2
                P_init, the initial 3D coordinates of points
        Output: M2, the optimized extrinsics of camera 1
                P2, the optimized 3D coordinates of points
                o1, the starting objective function value with the initial input
                o2, the ending objective function value after bundle adjustment
    
    ***
    Hints:
    (1) Use the scipy.optimize.minimize function to minimize the objective function, rodriguesResidual. 
        You can try different (method='..') in scipy.optimize.minimize for best results. 
    '''
    # print("lets ee if its ")
    obj_start = obj_end = 0
    # ----- TODO -----
    # YOUR CODE HERE
    N = p1.shape[0]
    t = M2_init[:,-1]
    R = M2_init[:,:3]
    r = invRodrigues(R)
    
    x = np.concatenate([P_init.flatten(), r.flatten(), t.flatten()])
    obj_start = get_residual_norm(x, K1, M1, p1, K2, p2)

    params = scipy.optimize.minimize(get_residual_norm, x, (K1, M1, p1, K2, p2))
    x_optimized = params.x

    P2 = x_optimized[:3*N]
    P2 = P2.reshape((N,3))
    
    t2 = x_optimized[-3:]
    t2 = t2.reshape((3,1))
    r2 = x_optimized[-6:-3]
    r2 = r2.reshape((3,1))

    R2 = rodrigues(r2)
    M2 = np.hstack((R2,t2))

    obj_end = get_residual_norm(x_optimized, K1, M1, p1, K2, p2)
    # raise NotImplementedError()
    return M2, P2, obj_start, obj_end

```

```python
# Visualization:
np.random.seed(0)
correspondence = np.load('data/some_corresp_noisy.npz') # Loading noisy correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1 = plt.imread('data/im1.png')
im2 = plt.imread('data/im2.png')

M=np.max([*im1.shape, *im2.shape])
F, inliners = ransacF(pts1, pts2, M)
pts1_inliners = pts1[inliners.squeeze(), :]
pts2_inliners = pts2[inliners.squeeze(), :]

M1 = np.hstack((np.identity(3), np.zeros(3)[:,np.newaxis]))
C1 = K1.dot(M1)
# print(F, inliners)
M2_init,C2, P_init = find_M2(F, pts1_inliners, pts2_inliners, intrinsics)

M2, P_final, obj_start, obj_end = bundleAdjustment(K1, M1, pts1_inliners, K2, M2_init, pts2_inliners, P_init)
print(f"Before {obj_start}, After {obj_end}")
# plot_3D_dual(P_init, P_final) ## comment to visualize, uncomment when submitting to gradescope

```

Output:
```
Best Error [52364.8016]
(3, 4) (3, 4)
(3, 4) (3, 4)
Before 238.92786884075895, After 11.348249171063372
File ~/.local/lib/python3.8/site-packages/matplotlib/backends/_backend_tk.py:75, in blit(photoimage, aggimage, offsets, bbox)
     73     bboxptr = (x1, x2, y1, y2)
     74 else:
---> 75     photoimage.blank()
     76     bboxptr = (0, width, 0, height)
     77 _tkagg.blit(
     78     photoimage.tk.interpaddr(), str(photoimage), dataptr, offsets, bboxptr)

File /usr/lib/python3.8/tkinter/__init__.py:4068, in PhotoImage.blank(self)
   4066 def blank(self):
   4067     """Display a transparent image."""
-> 4068     self.tk.call(self.name, 'blank')

TclError: invalid command name "pyimage10"
```

```python
# plot_3D_dual(P_init, P_final)
```

Output:
```
/home/harshad/learning/tutorial-env/CV/hw5/q3.ipynb Cell 16 in <cell line: 1>()
----> <a href='vscode-notebook-cell:/home/harshad/learning/tutorial-env/CV/hw5/q3.ipynb#X32sZmlsZQ%3D%3D?line=0'>1</a> plot_3D_dual(P_init, P_final)

/home/harshad/learning/tutorial-env/CV/hw5/q3.ipynb Cell 16 in plot_3D_dual(P_before, P_after)
     <a href='vscode-notebook-cell:/home/harshad/learning/tutorial-env/CV/hw5/q3.ipynb#X32sZmlsZQ%3D%3D?line=20'>21</a> ax.scatter(P_after[:,0], P_after[:,1], P_after[:,2], c='red')
     <a href='vscode-notebook-cell:/home/harshad/learning/tutorial-env/CV/hw5/q3.ipynb#X32sZmlsZQ%3D%3D?line=21'>22</a> while True:
---> <a href='vscode-notebook-cell:/home/harshad/learning/tutorial-env/CV/hw5/q3.ipynb#X32sZmlsZQ%3D%3D?line=22'>23</a>     x, y = plt.ginput(1, mouse_stop=2)[0]
     <a href='vscode-notebook-cell:/home/harshad/learning/tutorial-env/CV/hw5/q3.ipynb#X32sZmlsZQ%3D%3D?line=23'>24</a>     plt.draw()

IndexError: list index out of range
```

```python
a = np.zeros((3,3))
a
b = np.concatenate((a, np.ones((3,1))), axis=1)
b
```